_Main topics covered during today's session:_

Prior NBs:

1. **Counting values in a dictionary**

2. **Comparing dictionaries, and finding the highest/lowest value.**
    
    
This NB:

3. **Inverting a dictionary**


Following NB:

4. **Teaser -- Floating point values**

# Inverting a dictionary.

### This notebook was part of MT1 in the first iteration of CSE6040, when the MT's were a number of smaller and independent notebooks.

#### We will be doing Exercise 0, which is inverting a dictionary. What the means is to take the values of the dictionary and make them keys, and to take the keys and make them values, of a new dictionary.

# Problem 8

This problem has three (3) exercises worth a total of ten (10) points.

**Invert a dictionary.** Suppose a group of student TAs have been asked to monitor the Piazza forums for your CSE6040x course. The assignments are as follows.

```python
assignments = {"Monday" : ["Jeh", "Ben", "Chinmay", "Rachel"],
               "Tuesday" : ["Rachel", "Samuel", "Evan", "Raghav"],
               "Wednesday" : ["Evan", "Samuel", "Michael", "Chinmay"],
               "Thursday" : ["Jeh", "Michael", "Ben", "Evan"],
               "Friday" : ["Rachel", "Shishir", "Samuel", "Ben"],
               "Saturday" : ["Yandi", "Michael", "Chinmay", "Samuel"],
               "Sunday" : ["Evan", "Chinmay", "Jeh", "Rachel"]
               }
```

**Exercise 0** (5 points). Write a function, `invert_dict(d)`, that "inverts" a given dictionary `d`. The output should be another dictionary such that the keys are the TA names and their corresponding values should be a list which stores the days assigned to them.

For example, one of the keys of the returned dictionary should be `"Jeh"` with its corresponding list as `["Monday", "Thursday", "Sunday"]`.

```python
from collections import defaultdict
def invert_dict(d):
    ###
    ### YOUR CODE HERE
    ###

    # Larry's code, not sample solution
    # TA_dict = {}
    from collections import defaultdict
    TA_dict = defaultdict(list)

    for key, val in d.items():
        # print(key, val, sep='\n')
        for name in val:
            # print(key, name, sep='\n')
            # good template code for populating a dict,
            # while not using default dict
            # if not name in TA_dict.keys():
            #     TA_dict[name] = [key]
            # else:  # TA name is already a key in the dict
                # day_list = TA_dict[name]  # list of days
            TA_dict[name].append(key)
                # TA_dict[name] = day_list
            # print(TA_dict)
            # print()

#     print(TA_dict)

    return dict(TA_dict)

inv_dict = invert_dict(assignments)
```

<!-- grade_id: test_inverse, points: 5, grade: true, solution: false -->
```python
# Test cell: `test_inverse`

import random
TA_list = ["Rachel", "Yandi", "Ben", "Jeh", "Evan", "Chinmay", "Shishir", "Raghav", "Samuel", "Michael"]

# Your solution
inv_dict = invert_dict(assignments)

random_TA = random.sample(TA_list, 5)
for TA in random_TA:
    assigned_days = inv_dict[TA]
    for days in assigned_days:
        assert TA in assignments[days], "Incorrect inversion for TA {}".format(TA)

print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Exercise 2** (2 points): _Tracing the route of a pilot_.

Suppose a pilot flies from one city to the next. He makes 7 such trips, which are stored in the list shown in the next cell. The first entry of each tuple denotes the origin, and the second denotes the destination. Also assume that the next flight a pilot makes must originate from the destination of her previous flight. Your task in this exercise is to write a function that finds the route followed by the pilot, given her first port of origin.

First, let `segments` be an unordered list of segments that the pilot flew.

```python
segments = [("JFK", "DEN"), ("LAX", "ORD"), ("DEN", "SFO"), ("LAS", "LAX"), ("ORD", "ATL"), ("ATL", "JFK"), ("SFO", "LAS")]
```

Next, write a function `find_dest(segs, origin)` that returns the next destination of the pilot given one port of origin and an unordered list of flight segments. Example, if `segs == segments` as defined above, then for `'LAX'`, your function should return `'ORD'` because there is the tuple, `("LAX", "ORD")`, in `segs`.

You may assume that `origin` appears only once in the list of segments, `segs`.

```python
def find_dest(segs, origin):
    assert type(origin) is str
    assert type(segs) is list
    ###
    ### YOUR CODE HERE
    ###

    dest = ''

    for route in segs:
        print('tick')
        if route[0] == origin:
            return route[1]
    # return [d for o, d in segs if o == origin][0]

    return dest
    # route_d = dict(segs)
    # return route_d[origin]

find_dest(segments, "DEN")
```

Output:
```
tick
tick
tick
'SFO'
```

<!-- grade_id: test_destination, points: 2, grade: true, solution: false -->
```python
# Test flight destination
den = find_dest(segments, "DEN")
assert find_dest(segments, "DEN") == "SFO", "Wrong destination for DEN"
assert find_dest(segments, "LAX") == "ORD", "Wrong destination for LAX"
assert find_dest(segments, "ATL") == "JFK", "Wrong destination for ATL"
assert find_dest(segments, "JFK") == "DEN", "Wrong destination for JFK"
print("\n(Passed.)")
```

Output:
```
(Passed.)
```

**Exercise 3** (3 points). Now write a function `get_route(segs, origin)` that traces the complete itinerary of the pilot, given her first port of origin. Assume that the destination for a flight is the same as the origin for the next flight. The itinerary is completed when he arrives back at the starting port of origin. For example, if the starting port is `"JFK"`, your function should return the list: `["JFK", "DEN", "SFO", "LAS", "LAX", "ORD", "ATL"]`.

```python
def get_route(path, origin):
    ###
    ### YOUR CODE HERE
    ###

    # path = segs

    # route_list = [origin]
    # for dest in path:
    #     print('tock')
    #     route_list.append(find_dest(path,origin))
    #     origin = route_list[-1]
    # print(route_list)
    # route_list.pop()


    # return route_list
    route_d = dict(path)
    start = route_d[origin] #second starting point
    route_list = [origin] # pre-populate with first trip
    for _ in range(len(path)):
      print('tick')
    while start != origin: # captures from second trip on
      print('tock')
      route_list.append(start)
      start = route_d[start]
    return route_list

itinerary = get_route(segments, 'JFK')
print(itinerary)
```

Output:
```
tick
tick
tick
tick
tick
tick
tick
tock
tock
tock
tock
tock
tock
['JFK', 'DEN', 'SFO', 'LAS', 'LAX', 'ORD', 'ATL']
```

<!-- grade_id: test_route, points: 3, grade: true, solution: false -->
```python
# Test pilot route
ports = ["JFK", "DEN", "SFO", "LAS", "LAX", "ORD", "ATL"]
starting_port = random.sample(ports, 5)
for p in starting_port:
    itinerary = get_route(segments, p)
    assert itinerary[0] == p, "incorrect port of origin for the itinerary"
    for i, port in enumerate(itinerary[:-1]):
        dest = find_dest(segments, port)
        assert dest == itinerary[i+1], "incorrect itinerary"

print("\n(Passed.)")
```

Output:
```
tick
tick
tick
tick
tick
tick
tick
tock
tock
tock
tock
tock
tock
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tock
tock
tock
tock
tock
tock
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tock
tock
tock
tock
tock
tock
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tock
tock
tock
tock
tock
tock
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tock
tock
tock
tock
tock
tock
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick
tick

(Passed.)
```

**Fin!** You've reached the end of this problem. Don't forget to restart the
kernel and run the entire notebook from top-to-bottom to make sure you did
everything correctly. If that is working, try submitting this problem. (Recall
that you *must* submit and pass the autograder to get credit for your work!)